# Coding cross validation yourself

Now it is up to you to demonstrate your knowledge on the absenteeism dataset.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd


np.random.seed(40)

data = pd.read_csv("absent.csv",sep=';',index_col=0)
y = data['Absenteeism time in hours']
X = data.drop('Absenteeism time in hours',axis=1)
```

## Applying cross validation

Apply cross-validation to this regression, where you normalize all variables, and calculate the mean absolute error, mean squared error, and R-squared value (only return the test set value):

<!-- grade_id: apply_cv, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
classifier = LinearRegression()    
    
def apply_cv(X, y, nFold):    
    # Your ouput object should be a cross_validate() result
    from sklearn.pipeline import make_pipeline
    from sklearn.model_selection import cross_validate
    from sklearn.preprocessing import Normalizer
    from sklearn.model_selection import cross_val_predict
    from sklearn.metrics import mean_squared_error
    
    normalizer = Normalizer()
    pipeline = make_pipeline(normalizer, classifier)
    
    # Define metrics
    metrics = ['neg_mean_absolute_error', 'neg_mean_squared_error', 'r2']
    
    outcomes = cross_validate(pipeline, X, y, scoring=metrics, cv=nFold, return_train_score=False)
    ###
    ### YOUR CODE HERE
    ###
    
    return outcomes
```

Validate your result:

<!-- grade_id: check_apply_cv, points: 5, grade: true, solution: false -->
```python
assert np.allclose(apply_cv(X,y,10)['test_neg_mean_absolute_error'],[-5.44169706, -4.88533949, -6.03776489, -6.53316679, -6.97119581,
       -5.6121378 , -6.20582349, -5.11806841, -5.58925382, -9.20383222])


###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

## Cross validation with prediction

Calculate the root mean squared error of a prediction when it was in the test set:

<!-- grade_id: cell-2622132d0c9d2aed, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression
classifier = LinearRegression()    
    
def calculate_rmse(X, y, nFold):  
    rmse = 0
    
    from sklearn.pipeline import make_pipeline
    from sklearn.model_selection import cross_validate
    from sklearn.preprocessing import Normalizer
    from sklearn.model_selection import cross_val_predict
    from sklearn.metrics import mean_squared_error
    from math import sqrt
    
    normalizer = Normalizer()
    pipeline = make_pipeline(normalizer, classifier)
    
    prediction = cross_val_predict(pipeline, X, y, cv=nFold)
    rmse = sqrt(mean_squared_error(y, prediction))
    ###
    ### YOUR CODE HERE
    ###
    
    return rmse
```

Validate your result:

<!-- grade_id: None, points: 3, grade: true, solution: false -->
```python
assert np.allclose(calculate_rmse(X, y, 10),12.677377376954103)
assert np.allclose(calculate_rmse(X, y, 5),12.846765912010982)


###
### AUTOGRADER TEST - DO NOT REMOVE
###

```