# Exercise 1

An easy one to start with. Calculate the number of letters in your name:

<!-- grade_id: get_length_name, points: 0, grade: false, solution: true -->
```python
first_name = "jinglin"
last_name = "wong"

def get_length_name(first_name, last_name):
    length = 0
    length =len(first_name)+len(last_name)
    ###
    ### YOUR CODE HERE
    ###
    
    return length
```

Verify your answer (you typically get 1 answer for free, so you are aware of how the exercises are validated):

<!-- grade_id: get_length_name, points: 5, grade: true, solution: false -->
```python
assert get_length_name('Alexander','the Great') == 18

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

# Exercise 2

The second exercise focuses on arithmetic and arrays.
Find the sum of the elements of two arrays.

<!-- grade_id: get_array_sums, points: 0, grade: false, solution: true -->
```python
def get_array_sums(arr1, arr2):
    sum1 = sum(arr1)
    sum2 = sum(arr2)
    
    ###
    ### YOUR CODE HERE
    ###
    
    return sum1+sum2
```

Verify your answer:

<!-- grade_id: get_array_sums, points: 5, grade: true, solution: false -->
```python
assert get_array_sums([1, 2, 3, 4, 5],[6, 7, 8, 9, 10]) == 55

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

# Exercise 3

Calculate how many years you need to obtain a certain amount by starting with an initial amount and a given interest rate. Don't forget: we use compound interest: you get interest on your interest.

<!-- grade_id: calculate_interest, points: 0, grade: false, solution: true -->
```python
import numpy as np
def calculate_interest(amount, target, interest_rate):
    number_of_years = np.log((target/amount))/np.log(1+interest_rate)+1
    
    ###
    ### YOUR CODE HERE
    ###
    
    return int(number_of_years)
```

Verify your answer:

<!-- grade_id: calculate_interest, points: 8, grade: true, solution: false -->
```python
assert calculate_interest(1000,2000,0.02) == 36

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

# Exercise 4

Remove the duplicates from a list.

<!-- grade_id: remove_duplicates, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(input_list):
    output_list = []
    for i in input_list:
        if i not in output_list:
            output_list.append(i)
    ###
    ### YOUR CODE HERE
    ###
    
    return output_list
```

Verify your answer:

<!-- grade_id: remove_duplicates, points: 4, grade: true, solution: false -->
```python
assert remove_duplicates(["Johannes", "Jamal", "Jamal", "Johannes", "Galina"]) == ['Johannes', 'Jamal', 'Galina']

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

# Exercise 5

Find the total number of points taught by a particular individual. The input is two dictionaries: one containing the courses and their organisers, and one containing the courses and their points.

<!-- grade_id: cell-2aa20234cf44249b, points: 0, grade: false, solution: true -->
```python
def find_points_for_organizer(organizer, courses, points):
    total = 0
    for course, man in courses.items():
        if man == organizer:
            total += points[course]
    ###
    ### YOUR CODE HERE
    ###
    
    return total
```

Verify your answer:

<!-- grade_id: find_points_for-organizer, points: 6, grade: true, solution: false -->
```python
courses = {"Analytics": "Johannes", "Mega analytics" : "Johannes", "Basic analytics" : "John"}
points = {"Analytics": 15, "Mega analytics": 10, "Basic analytics":5}

assert find_points_for_organizer('Johannes', courses, points) == 25

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```