# Exercises with libraries

Before you begin, load the wine dataset, then perform the tasks outlined below.

```python
from sklearn import datasets as ds
import pandas as pd
import numpy as np

dataset = ds.load_wine()
dataset.keys()
```

Output:
```
dict_keys(['data', 'target', 'target_names', 'DESCR', 'feature_names'])
```

First, retrieve all the names of the independent variables:

<!-- grade_id: get_variable_names, points: 0, grade: false, solution: true -->
```python
def get_variable_names(dataset):
    variables = dataset['feature_names']
    ###
    ### YOUR CODE HERE
    ###
    
    return variables
```

```python
get_variable_names(dataset)
```

Output:
```
['alcohol',
 'malic_acid',
 'ash',
 'alcalinity_of_ash',
 'magnesium',
 'total_phenols',
 'flavanoids',
 'nonflavanoid_phenols',
 'proanthocyanins',
 'color_intensity',
 'hue',
 'od280/od315_of_diluted_wines',
 'proline']
```

```python
df = pd.DataFrame(data=dataset['data'], columns=dataset['feature_names'])
```

Your answer will be verified below (no need for you to do anything).

<!-- grade_id: get_variables_names, points: 2, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
df.head()
```

Output:
```
alcohol  malic_acid   ash  alcalinity_of_ash  magnesium  total_phenols  \
0    14.23        1.71  2.43               15.6      127.0           2.80   
1    13.20        1.78  2.14               11.2      100.0           2.65   
2    13.16        2.36  2.67               18.6      101.0           2.80   
3    14.37        1.95  2.50               16.8      113.0           3.85   
4    13.24        2.59  2.87               21.0      118.0           2.80   

   flavanoids  nonflavanoid_phenols  proanthocyanins  color_intensity   hue  \
0        3.06                  0.28             2.29             5.64  1.04   
1        2.76                  0.26             1.28             4.38  1.05   
2        3.24                  0.30             2.81             5.68  1.03   
3        3.49                  0.24             2.18             7.80  0.86   
4        2.69                  0.39             1.82             4.32  1.04   

   od280/od315_of_diluted_wines  proline  
0                     
```

Now, create a function that calculates the mean and median of a given variable in the dataset (enter the variable's name as a string):

<!-- grade_id: calculate_mean_median, points: 0, grade: false, solution: true -->
```python
def calculate_mean_median(variable):
    mean = df[variable].mean()
    median = df[variable].median()
    
    ###
    ### YOUR CODE HERE
    ###
    
    return mean, median
```

Finally, verify your answer (note: the function allclose is used to avoid any difference in rounding):

<!-- grade_id: calculate_mean_median, points: 4, grade: true, solution: false -->
```python
assert np.allclose(calculate_mean_median('flavanoids'),(2.0292696629213474, 2.135))

###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```